# Week1 Assignment - Testing OKPy and the autograder.

For this assignment, you will load the iris flowers dataset from sklearn and predict flower species based on their petal and sepal sizes.

```python
"""This is a docstring!"""
```

Output:
```
'This is a docstring!'
```

$\beta=5$

```python
alpha = 4
```

```python
%who
```

Output:
```
alpha
```

```python
var_list = %who_ls
```

```python
var_list
```

Output:
```
['alpha']
```

## Part A. Load the Iris dataset.

Sklearn comes with prepared datasets that you can load. To load the iris dataset, run:
```python
from sklearn import datasets
iris = datasets.load_iris()
```

```python
from sklearn import datasets
iris = datasets.load_iris()
```

Output:
```
/srv/conda/envs/notebook/lib/python3.7/importlib/_bootstrap.py:219: RuntimeWarning: numpy.ufunc size changed, may indicate binary incompatibility. Expected 192 from C header, got 216 from PyObject
  return f(*args, **kwds)
```

```python
iris.target_names
```

Output:
```
array(['setosa', 'versicolor', 'virginica'], dtype='<U10')
```

```python
iris.target
```

Output:
```
array([0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0,
       0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0, 0,
       0, 0, 0, 0, 0, 0, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1,
       1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1,
       1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 1, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2,
       2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2,
       2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2, 2])
```

<!-- grade_id: cell-c1e7f79ad332b8f9, points: 0, grade: false, solution: true -->
```python
from sklearn import datasets

# Load the iris dataset
# YOUR CODE HERE
raise NotImplementedError()
```

Output:
```
NotImplementedError                       Traceback (most recent call last)
<ipython-input-9-eb43bd51aa77> in <module>
      3 # Load the iris dataset
      4 # YOUR CODE HERE
----> 5 raise NotImplementedError()

NotImplementedError:
```

<!-- grade_id: cell-fd50146eeead0c65, points: 1, grade: true, solution: false -->
```python
assert (iris['data'] == datasets.load_iris()['data']).all()
```

## Part B. Use an SVM to predict iris species.

To setup the data and train an SVM, run the following:
```python
from sklearn.svm import SVC

X = iris.data
y = iris.target

classifier = SVC(kernel="rbf")
classifier.fit(X, y)
predictions = classifier.predict(X)
```

```python
%who
```

Output:
```
SVC	 X	 alpha	 classifier	 datasets	 iris	 predictions	 var_list	 y
```

```python
from sklearn.svm import SVC

X = iris.data
y = iris.target

X_train, X_test, y_train, y_test = train_test_split(X, y, test_size=0.33, random_state=42)

classifier = SVC()
```

```python
X.shape
```

Output:
```
(150, 4)
```

```python
X_train.shape
```

Output:
```
(100, 4)
```

```python
X_test.shape
```

Output:
```
(50, 4)
```

```python
y.shape
```

Output:
```
(150,)
```

```python
y_train.shape
```

Output:
```
(100,)
```

```python
y_test.shape
```

Output:
```
(50,)
```

```python
classifier.fit(X_train, y_train)

predictions = classifier.predict(X_test)
```

Output:
```
/srv/conda/envs/notebook/lib/python3.7/site-packages/sklearn/svm/base.py:193: FutureWarning: The default value of gamma will change from 'auto' to 'scale' in version 0.22 to account better for unscaled features. Set gamma explicitly to 'auto' or 'scale' to avoid this warning.
  "avoid this warning.", FutureWarning)
```

```python
classifier.score(X_test,y_test)
```

Output:
```
1.0
```

X = features, y = predictions
>Will only use random state when you need some element of randomness (shuffling) to the data (ex: test_train_split)
>We just scored the model using classifer.score()

- Overfitting (Variance): tailored to a specific dataset (does not generalize well), misses the point, memorizes the data
- Underfitting (Bias): accuracy could be better, too general 

Scikit-learn workflow
1. Import
2. Instantiate (a class)
  - Estimators:
    - Classifiers
    - Regressors
3. Fit (fit())
4. Predict (predict())
5. Score (score())

```python
classifier.score(X,y)
```

Output:
```
0.9866666666666667
```

##### score() makes its own predictions, and then scores them
##### most classification models uses scoring as a metric

```python
predictions
```

```python
from sklearn.model_selection import train_test_split
```

```python
train_test_split?
```

Output:
```
[0;31mSignature:[0m [0mtrain_test_split[0m[0;34m([0m[0;34m*[0m[0marrays[0m[0;34m,[0m [0;34m**[0m[0moptions[0m[0;34m)[0m[0;34m[0m[0;34m[0m[0m
[0;31mDocstring:[0m
Split arrays or matrices into random train and test subsets

Quick utility that wraps input validation and
``next(ShuffleSplit().split(X, y))`` and application to input data
into a single call for splitting (and optionally subsampling) data in a
oneliner.

Read more in the :ref:`User Guide <cross_validation>`.

Parameters
----------
*arrays : sequence of indexables with same length / shape[0]
    Allowed inputs are lists, numpy arrays, scipy-sparse
    matrices or pandas dataframes.

test_size : float, int or None, optional (default=None)
    If float, should be between 0.0 and 1.0 and represent the proportion
    of the dataset to include in the test split. If int, represents the
    absolute number of test samples. If None, the value is set to the
    complement of the train size. If ``train_size`` is
```

<!-- grade_id: cell-bc8ce57a60e5fd7a, points: 0, grade: false, solution: true -->
```python
from sklearn.svm import SVC

# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-165c467d2a6e7dde, points: 1, grade: true, solution: false -->
```python
assert X.shape == (150, 4)
assert len(y) == 150
```